```python
# Initialize Otter
import otter
grader = otter.Notebook("25100072_2.ipynb")
```

# **Assignment 1 (Part 2)**

<h3>Assignment Instructions</h3>


* **How to install 'otter'**: Run `pip install otter-grader` in your Anaconda Command Prompt
* **Otter**: It is an autograder that we will be using for grading your notebooks.
* **grader.check('qn')**: This command runs test cases for nth question `qn` provided to you and displays the result. But these are not the only test cases, these are provided just for your basic testing. There will be some *hidden test cases* on which your answers will be tested as well.
* You are **not** allowed to edit any pre-defined variables and as per the instructions for every question you have to assign your answers to these variables.
* Write your answers in place of `...`


<h3>Submission Instructions</h3>


* There are two parts of this assignment. You should name both notebooks as **YourRollNumber_PartNumber.ipynb**. (e.g. this part should be named as *`2XXXXXXX_2.ipynb`*)
* You must submit a zip file containg just two notebooks and it must not has any folder or any other file. 
* The name of the zip file must be your `RollNumber_PA1`


Please run the cell below to import all the necessary libraries that you would require for this part

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

<h3>Context about the datasets</h3>

Given the rise in obesity we are aware that there is a huge chance that a person may be affected by the heart disease. In this assignment we are going to explore a dataset which contains information about individuals heart rate and oxygen saturation. In this part we are going to analyze the data and generate sufficent results from it.

Note: The output column is binary (0 or 1) and signifies whether the individual has had a heart attack (1 if they have had a heart attack).

**Question 0:** Please load the two dataset files inside the variables `heartDf` and `o2SaturationDf` respectively

```python
heartDf = pd.read_csv('heart.csv')
o2SaturationDf = pd.read_csv('o2saturation.csv')
```

```python
grader.check("Q0")
```

Output:
```
Q0 results: All test cases passed!
```

**Question 1:** Find the shapes of both datasets and assign them to `heartDfShape` and `o2DfShape` respectively

```python
heartDfShape = heartDf.shape
o2DfShape = o2SaturationDf.shape
```

```python
grader.check("Q1")
```

Output:
```
Q1 results: All test cases passed!
```

**Question 2A:** Use a pandas function to display all characteristics (including min, max, quartiles, count, mean) of all numerical attributes of both dataframes. Store these in the variables `heartdf_description` & `o2df_description` respectively. 

```python
heartdf_description = heartDf.describe()
o2df_description = o2SaturationDf.describe()
```

```python
grader.check("Q2a")
```

Output:
```
Q2a results: All test cases passed!
```

**Question 2B:** Use a pandas function to display the number of occurances of all unique values of sex in the dataset. Store these in the variable `gender_distribution`. 

```python
gender_distribution = heartDf['sex'].value_counts() 
```

```python
grader.check("Q2B")
```

Output:
```
Q2B results: All test cases passed!
```

**Question 3A:** Prepare the two dataframes for merging by making the following column changes to both the dataframes. 
<h5>For Heart Dataset</h5>
<ul>
<li>exng -> exercise induced angina</li>
<li>caa -> number of major vessels</li>
<li>cp -> chest pain type</li>
<li>trtbps -> resting blood pressure</li>
<li>chol -> cholestrol levels</li>
<li>restecg -> resting electrocardiographic </li> 
<li>thalachh-> maximum heart rate</li>
<li>fbs-> fasting blood sugar > 120mg/dl</li>
<li>oldpeak -> old peak </li>
</ul>
<h5>For O2 Saturation Dataset </h5>
<ul>
<li>sat -> O2 saturation</li>
</ul>


```python
heartDf.rename(columns={'exng':'exercise induced angina','caa':'number of major vessels','cp':'chest pain type','trtbps':'resting blood pressure','chol':'cholesterol levels','restecg':'resting electrocardiographic','thalachh':'maximum heart rate','fbs':'fasting blood sugar > 120mg/dl','oldpeak':'old peak'}, inplace=True) 
o2SaturationDf.rename(columns={'sat':'O2 saturation'}, inplace=True)
```

```python
grader.check("Q3A")
```

Output:
```
Q3A results: All test cases passed!
```

```python
o2SaturationDf.head(5)

```

Output:
```
O2 saturation
0           98.6
1           98.6
2           98.6
3           98.1
4           97.5
```

```python
heartDf.head(5)
```

Output:
```
age  sex  chest pain type  resting blood pressure  cholesterol levels  \
0   63    1                3                   145.0               233.0   
1   37    1                2                   130.0               250.0   
2   41    0                1                   130.0               204.0   
3   56    1                1                   120.0               236.0   
4   57    0                0                   120.0               354.0   

   fasting blood sugar > 120mg/dl  resting electrocardiographic  \
0                               1                             0   
1                               0                             1   
2                               0                             0   
3                               0                             1   
4                               0                             1   

   maximum heart rate  exercise induced angina  old peak  slp  \
0               150.0                        0       2.3    0   
1            
```

```python
heartDf.columns
```

Output:
```
Index(['age', 'sex', 'chest pain type', 'resting blood pressure',
       'cholesterol levels', 'fasting blood sugar > 120mg/dl',
       'resting electrocardiographic', 'maximum heart rate',
       'exercise induced angina', 'old peak', 'slp', 'number of major vessels',
       'thall', 'output'],
      dtype='object')
```

**Question 3B:** Peform a merge operation so that we can better compare our results and reach an appropriate conclusion. 

```python
mergedDf = pd.merge(heartDf, o2SaturationDf, left_index=True, right_index=True, how='outer')
```

```python
grader.check("Q3B")
```

Output:
```
Q3B results: All test cases passed!
```

```python

print(mergedDf.dtypes)

```

Output:
```
age                                 int64
sex                                 int64
chest pain type                     int64
resting blood pressure            float64
cholesterol levels                float64
fasting blood sugar > 120mg/dl      int64
resting electrocardiographic        int64
maximum heart rate                float64
exercise induced angina             int64
old peak                          float64
slp                                 int64
number of major vessels             int64
thall                               int64
output                              int64
O2 saturation                     float64
dtype: object
```

**Question 4:** Categorize the individuals based on their age group using the following groupings in the corresponding variable names with `loc`:
<ul>
<li>Below 40 years of age</li>
<li>In the range of 40-50 years of age</li>
<li>In the range of 50-60 years of age</li>
<li>In the range of 60-70 years of age</li>
<li>70 or above years of age</li>
</ul>

<b>Note: </b> Only the lower bound is inclusive in the above age group range. You are to store the results in the variables given in the cell below.

```python
belowForty = mergedDf.loc[mergedDf['age'] < 40]
ageBtw40_50 = mergedDf.loc[(mergedDf['age'] >= 40) & (mergedDf['age'] < 50)]
ageBtw50_60 = mergedDf.loc[(mergedDf['age'] >= 50) & (mergedDf['age'] < 60)]
ageBtw60_70 = mergedDf.loc[(mergedDf['age'] >= 60) & (mergedDf['age'] < 70)]
aboveSeventy = mergedDf.loc[(mergedDf['age'] >= 70) ]
```

```python
grader.check("Q4")
```

Output:
```
Q4 results: All test cases passed!
```

**Question 5A:**  Group the data by 'sex' and find the average age for each group. <br/>

Store the result in the variable `mean_ages_by_sex`.

```python
mean_ages_by_sex = mergedDf.groupby('sex')['age'].mean()
```

```python
grader.check("Q5A")
```

Output:
```
Q5A results: All test cases passed!
```

**Question 5B:**  Extract rows for patients whose maximum heart rate is within 10% of the mean heart rate across all patients. <br/>

Store the result in the variable `patients_with_normal_heartrate`.

```python
mean_hr = mergedDf['maximum heart rate'].mean()

lower_bound = mean_hr * 0.9  # 10% below the mean
upper_bound = mean_hr * 1.1  # 10% above the mean

patients_with_normal_heartrate = mergedDf.loc[(mergedDf['maximum heart rate'] >= lower_bound) & (mergedDf['maximum heart rate'] <= upper_bound)]
```

```python
grader.check("Q5B")
```

Output:
```
Q5B results: All test cases passed!
```

**Question 5C:**  For patients younger than 50, find the average cholesterol levels for those with a maximum heart rate above 160. <br/>

Store the result in the variable `mean_cholesterol_youngerPatients`.

```python
df_filtered = pd.concat([belowForty,ageBtw40_50])
mean_cholesterol_youngerPatients = df_filtered.loc[(df_filtered['maximum heart rate'] > 160)].mean()
mean_cholesterol_youngerPatients = mean_cholesterol_youngerPatients[4]
```

Output:
```
C:\Users\sehar\AppData\Local\Temp\ipykernel_26444\4028516992.py:3: FutureWarning: Series.__getitem__ treating keys as positions is deprecated. In a future version, integer keys will always be treated as labels (consistent with DataFrame behavior). To access a value by position, use `ser.iloc[pos]`
  mean_cholesterol_youngerPatients = mean_cholesterol_youngerPatients[4]
```

```python
print(mean_cholesterol_youngerPatients)
```

Output:
```
232.3653846153846
```

```python
grader.check("Q5C")
```

Output:
```
Q5C results: All test cases passed!
```

**Question 5D:** Define a custom function that classifies a patient as `high risk` if their cholesterol level is above 240 and their `maximum heart rate` is below 120. Apply this function to each group of patients with different `chest pain type`s and count how many patients meet these high-risk conditions for each chest pain type. You will be using `.apply` for this task. <br/>

Store the result in the variable `high_risk_counts`. 

```python
def high_risk_filter(row):
    if row['cholesterol levels'] > 240 and row['maximum heart rate'] < 120:
        return True
    else:
        return False
    
```

```python
grouped_by_chest_pain = mergedDf.groupby('chest pain type')

```

```python
high_risk_counts = grouped_by_chest_pain.apply(lambda group: group.apply(high_risk_filter, axis=1).sum())

```

```python
grader.check("Q5D")
```

Output:
```
Q5D results: All test cases passed!
```

**Question 6:**  Is the average resting blood pressure of individuals in the (Above 70 age group) greater than the resting blood pressure of inviduals that are below 40 years old? <br/>
Store the result in `meanOfAboveSeventy` & `meanOfBelowForty` variables. <br/> Also the variable `ans` must contain a string such as a 'Yes' or 'No' as the test cases check for it. <br>
<b>Note: </b>You are not allowed to use the mean() method but rather calculate the mean mathematically.

```python
meanOfAboveSeventy = aboveSeventy['resting blood pressure'].sum() / len(aboveSeventy)
meanOfBelowForty = belowForty['resting blood pressure'].sum() / len(belowForty)


if meanOfAboveSeventy > meanOfBelowForty:
    q6ans = "Yes"
else:
    q6ans = "No"
```

```python
grader.check("Q6")
```

Output:
```
Q6 results: All test cases passed!
```

**Question 7:** What is the most occuring value for oxygen saturation in individuals that fall into the age group of 50-60 ? <br>
Store the value in the variable `frequentO2Saturation`.

```python
frequentO2Saturation = mergedDf['O2 saturation'].mode().values[0]
print(frequentO2Saturation)

```

Output:
```
97.5
```

```python
grader.check("Q7")
```

Output:
```
Q7 results: All test cases passed!
```

**Question 8:** Find the most prevalent type of chest pain for all age groups mentioned above and store it in the corresponding variable. Also, find the most prevalent type of chest pain across all the individuals and store your result in `mostObservedChestPain`.

```python
belowFortyChestPain = belowForty['chest pain type'].mode().values[0]
ageBtw40_50ChestPain = ageBtw40_50['chest pain type'].mode().values[0]
ageBtw50_60ChestPain = ageBtw50_60['chest pain type'].mode().values[0]
ageBtw60_70ChestPain = ageBtw60_70['chest pain type'].mode().values[0]
aboveSeventyChestPain = aboveSeventy['chest pain type'].mode().values[0]
mostObservedChestPain = mergedDf['chest pain type'].mode().values[0]
```

```python
grader.check("Q8")
```

Output:
```
Q8 results: All test cases passed!
```

**Question 9:** Calculate the <b>median</b> cholestrol levels for all the age groups.<br>
<b>Note: </b> Here the variable `listOfDf` contains all the age groups that you have previously calculated in question <b>6</b>.

```python
listOfDf = [belowForty,ageBtw40_50,ageBtw50_60,ageBtw60_70,aboveSeventy]
medianOfAllCholLevels = [df['cholesterol levels'].median() for df in listOfDf]
```

```python
grader.check("Q9")
```

Output:
```
Q9 results: All test cases passed!
```

**Question 10:**  Using a pivot table, show the gender distribution across dfferent chest pain types. Store the result in `gender_count`.

```python
gender_count = mergedDf.pivot_table(index='chest pain type', columns='sex', aggfunc='size', fill_value=0)
```

```python
grader.check("Q10")
```

Output:
```
Q10 results: All test cases passed!
```

**Question 11:**  Calculate the percentage of males (1) and females (0) who are diagnosed with exercise induced angina. Round the result to three decimal places and store the result in `male_prop` and `female_prop` respectively<br/>
<b>Note:</b> 1 means they have exercise induced angina whereas 0 means they do not. 

```python
# Calculate total counts of males and females
total_males = mergedDf[mergedDf['sex'] == 1].shape[0]
total_females = mergedDf[mergedDf['sex'] == 0].shape[0]

# Calculate counts of males and females diagnosed with exercise induced angina (1)
males_with_angina = mergedDf[(mergedDf['sex'] == 1) & (mergedDf['exercise induced angina'] == 1)].shape[0]
females_with_angina = mergedDf[(mergedDf['sex'] == 0) & (mergedDf['exercise induced angina'] == 1)].shape[0]

# Calculate the percentages
male_prop = round((males_with_angina / total_males) * 100, 3) if total_males > 0 else 0
female_prop = round((females_with_angina / total_females) * 100, 3) if total_females > 0 else 0
```

```python
grader.check("Q11")
```

Output:
```
Q11 results: All test cases passed!
```

**Question 12:**  What is the maximum heart rate across the different age groups ? 
Store the max heart rates obtained in the variable `heartRates`.

```python
heartRates =  [df['maximum heart rate'].max() for df in listOfDf]
heartRates
```

Output:
```
[202.0, 194.0, 195.0, 179.0, 162.0]
```

```python
grader.check("Q12")
```

Output:
```
Q12 results: All test cases passed!
```

```python
mergedDf.columns
```

Output:
```
Index(['age', 'sex', 'chest pain type', 'resting blood pressure',
       'cholesterol levels', 'fasting blood sugar > 120mg/dl',
       'resting electrocardiographic', 'maximum heart rate',
       'exercise induced angina', 'old peak', 'slp', 'number of major vessels',
       'thall', 'output', 'O2 saturation'],
      dtype='object')
```

```python
mergedDf
```

Output:
```
age  sex  chest pain type  resting blood pressure  cholesterol levels  \
0     63    1                3                   145.0               233.0   
1     37    1                2                   130.0               250.0   
2     41    0                1                   130.0               204.0   
3     56    1                1                   120.0               236.0   
4     57    0                0                   120.0               354.0   
..   ...  ...              ...                     ...                 ...   
304   57    1                0                   123.0                 NaN   
305   58    0                1                     NaN               235.0   
306   58    0                1                     NaN               235.0   
307   58    0                1                     NaN               189.0   
308   58    0                1                   150.0               235.0   

     fasting blood sugar > 120mg/dl  resting electrocardiographic  \
```

**Question 13A:**  How many people have cholestrol values greater than the threshold value given below ? Store it in the variable `numberOfPeople`. 
<br>
<ul>
<li>Cholestrol Threshold Level: 300</li>
</ul>

Also calculate the probability of these people getting a heart attack and store the result, rounded off to three decimal places, in `probability`.


```python
numberOfPeople = mergedDf.loc[(mergedDf['cholesterol levels'] > 300)].count().values[0]
heart_attack_cases = mergedDf.loc[(mergedDf['cholesterol levels'] > 300) & (mergedDf['output'] == 1)]

probability = round(heart_attack_cases.shape[0] / numberOfPeople, 3) if numberOfPeople > 0 else 0
```

```python
print(probability)
```

Output:
```
0.535
```

```python
grader.check("Q13A")
```

Output:
```
Q13A results: All test cases passed!
```

```python
mergedDf.head(5)
```

Output:
```
age  sex  chest pain type  resting blood pressure  cholesterol levels  \
0   63    1                3                   145.0               233.0   
1   37    1                2                   130.0               250.0   
2   41    0                1                   130.0               204.0   
3   56    1                1                   120.0               236.0   
4   57    0                0                   120.0               354.0   

   fasting blood sugar > 120mg/dl  resting electrocardiographic  \
0                               1                             0   
1                               0                             1   
2                               0                             0   
3                               0                             1   
4                               0                             1   

   maximum heart rate  exercise induced angina  old peak  slp  \
0               150.0                        0       2.3    0   
1            
```

<!-- BEGIN QUESTION -->

**Question 13B:**  Do you think people with such cholesetrol levels have a higher probability of getting an heart attack than those with lower cholesterol levels? Justify this claim by providing a valid reason.

```python
# Probability of people getting heart attack with Higher Cholestrol Levels
numberOfPeople = mergedDf.loc[(mergedDf['cholesterol levels'] > 300)].count().values[0]
heart_attack_cases = mergedDf.loc[(mergedDf['cholesterol levels'] > 300) & (mergedDf['output'] == 1)]

probability = round(heart_attack_cases.shape[0] / numberOfPeople, 3) if numberOfPeople > 0 else 0

# Probability of people getting heart attack with Lower Cholestrol Levels

numberOfPeople_lower = mergedDf.loc[(mergedDf['cholesterol levels'] < 300)].count().values[0]
heart_attack_cases_lower = mergedDf.loc[(mergedDf['cholesterol levels'] < 300) & (mergedDf['output'] == 1)]

probability_lower = round(heart_attack_cases_lower.shape[0] / numberOfPeople_lower, 3) if numberOfPeople_lower > 0 else 0
```

```python
print(f'High Cholestrol: {probability}, Lower Cholestrol {probability_lower}')
```

Output:
```
High Cholestrol: 0.535, Lower Cholestrol 0.538
```

From the above calculation we can see that both groups are equally likely to have a heart attack. People with lower cholestrol levels have slightly higher chances than people with higher cholestrol.

<!-- END QUESTION -->

**This marks the end of Assignment 1 Part 2 :) <br> Hooray! You made it to the end now give your self a pat on the back!**

Tests for Q0 (points: 1)
```
>>> assert heartDf.shape == (309, 14)
>>> assert o2SaturationDf.shape == (309, 1)
```


Tests for Q1 (points: 1)
```
>>> assert heartDfShape == (309, 14)
>>> assert o2DfShape == (309, 1)
```


Tests for Q10 (points: 3)
```
>>> assert gender_count.shape == (4, 2)
>>> assert gender_count[0].sum() == 101
>>> assert gender_count[1].sum() == 208
```


Tests for Q11 (points: 4)
```
>>> assert male_prop == 37.019
>>> assert female_prop == 21.782
```


Tests for Q12 (points: 3)
```
>>> assert heartRates == [202.0, 194.0, 195.0, 179.0, 162.0]
```


Tests for Q13A (points: 3)
```
>>> assert numberOfPeople == 43
>>> assert probability == 0.535
```


Tests for Q2B (points: 2)
```
>>> assert gender_distribution[1] == 208
>>> assert gender_distribution[0] == 101
```


Tests for Q2a (points: 2)
```
>>> assert heartdf_description.shape == (8, 14)
>>> assert o2df_description.shape == (8, 1)
```


Tests for Q3A (points: 3)
```
>>> assert 'exercise induced angina' in heartDf.columns
>>> assert 'number of major vessels' in heartDf.columns
>>> assert 'chest pain type' in heartDf.columns
>>> assert 'fasting blood sugar > 120mg/dl' in heartDf.columns
>>> assert 'resting blood pressure' in heartDf.columns
>>> assert 'old peak' in heartDf.columns
>>> assert 'cholesterol levels' in heartDf.columns
>>> assert 'resting electrocardiographic' in heartDf.columns
>>> assert 'maximum heart rate' in heartDf.columns
>>> assert 'O2 saturation' in o2SaturationDf.columns
```


Tests for Q3B (points: 2)
```
>>> assert mergedDf.shape == (309, 15)
>>> assert 'O2 saturation' in mergedDf.columns
```


Tests for Q4 (points: 5)
```
>>> assert belowForty.shape == (16, 15)
>>> assert ageBtw40_50.shape == (72, 15)
>>> assert ageBtw50_60.shape == (131, 15)
>>> assert ageBtw60_70.shape == (80, 15)
>>> assert aboveSeventy.shape == (10, 15)
```


Tests for Q5A (points: 2)
```
>>> assert mean_ages_by_sex[0] == 55.79207920792079
>>> assert mean_ages_by_sex[1] == 53.77403846153846
```


Tests for Q5B (points: 3)
```
>>> assert patients_with_normal_heartrate.shape == (143, 15)
```


Tests for Q5C (points: 2)
```
>>> assert mean_cholesterol_youngerPatients == 232.3653846153846
```


Tests for Q5D (points: 3)
```
>>> assert high_risk_counts[0] == 15
>>> assert high_risk_counts[1] == 1
>>> assert high_risk_counts[2] == 2
>>> assert high_risk_counts[3] == 0
```


Tests for Q6 (points: 4)
```
>>> assert q6ans.strip().lower() == 'yes'
>>> assert meanOfAboveSeventy == 135.8
>>> assert meanOfBelowForty == 125.5
```


Tests for Q7 (points: 2)
```
>>> assert frequentO2Saturation == 97.5
```


Tests for Q8 (points: 5)
```
>>> assert belowFortyChestPain == 2
>>> assert ageBtw40_50ChestPain == 0
>>> assert ageBtw50_60ChestPain == 0
>>> assert ageBtw60_70ChestPain == 0
>>> assert aboveSeventyChestPain == 0
>>> assert mostObservedChestPain == 0
```


Tests for Q9 (points: 4)
```
>>> assert medianOfAllCholLevels == [207.0, 235.5, 239.0, 252.5, 267.0]
```
