```python
# Initialize Otter
import otter
grader = otter.Notebook("25100072_1.ipynb")
```

# **Assignment 1 (Part 1)**

<h3>Assignment Instructions</h3>


* **How to install 'otter'**: Run `pip install otter-grader` in your Anaconda Command Prompt
* **Otter**: It is an autograder that we will be using for grading your notebooks.
* **Initialize Otter**: Run the cell above to initialize otter. Replace `2XXXXXX_1.ipynb` with your updated filename inside the grader variable above. 
* **grader.check('qn')**: This command runs test cases for nth question `qn` provided to you and displays the result. But these are not the only test cases, these are provided just for your basic testing. There will be some *hidden test cases* on which your answers will be tested as well.
* You are **not** allowed to edit any pre-defined variables and as per the instructions for every question you have to assign your answers to these variables.
* Write your answers in place of `...`


<h3>Submission Instructions</h3>


* There are two parts of this assignment. You should name both notebooks as **YourRollNumber_PartNumber.ipynb**. (e.g. this part should be named as *`2XXXXXXX_1.ipynb`*)
* You must submit a zip file containg just two notebooks and it must not has any folder or any other file. 
* The name of the zip file must be your `RollNumber_PA1`


Please run the cell below to import all the necessary libraries that you would require for this part

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

<h3>Context about the dataset</h3>


The titanic dataset is a record of passengers aboard the Titanic and their fates (whether or not they survived the incident). This includes socio-economic class divide between passengers, their ages, names, fares and even their passenger IDs. The principal source for data about Titanic passengers is the Encyclopedia Titanica. The datasets used here were begun by a variety of researchers. One of the original sources is Eaton & Haas (1994) Titanic: Triumph and Tragedy, Patrick Stephens Ltd, which includes a passenger list created by many researchers and edited by Michael A. Findlay.

As we explore this data from titanic, it is your job to answer some pressing questions regarding the various age groups aboard, their economic statuses and survival rates in each gender. Following are some details about the variable to help you understand the data better. <br>

<h4>Variable Descriptions</h4>

* Pclass ->   Passenger Class  (1 = 1st; 2 = 2nd; 3 = 3rd) <br>
* Survived ->  Did the passenger survive? (0 = No; 1 = Yes) <br>
* Name ->  name <br>
* Sex ->  sex  <br>
* Age ->  age  <br>
* SibSp ->  Number of Siblings/Spouses Aboard <br>
* Parch ->  Number of Parents/Children Aboard <br>
* Ticket ->   Ticket Number <br>
* Fare ->   Passenger Fare (British pound) <br>
* Cabin ->   Cabin <br>
* Embarked ->   Port of Embarkation (C = Cherbourg; Q = Queenstown; S = Southampton) <br>

**Question 0:** Load the titanic dataset file into the variable `titanic_df`

```python
titanic_df = pd.read_csv('titanic.csv')
```

```python
grader.check("Q0")
```

Output:
```
Q0 results: All test cases passed!
```

```python
titanic_df
```

Output:
```
Age Cabin Embarked     Fare  \
0    22.0   NaN        S   7.2500   
1    38.0   C85        C  71.2833   
2    26.0   NaN        S   7.9250   
3    35.0  C123        S  53.1000   
4    35.0   NaN        S   8.0500   
..    ...   ...      ...      ...   
886  27.0   NaN        S  13.0000   
887  19.0   B42        S  30.0000   
888  22.0   NaN        S  23.4500   
889  26.0  C148        C  30.0000   
890  32.0   NaN        Q   7.7500   

                                                  Name  Parch  PassengerId  \
0                              Braund, Mr. Owen Harris      0            1   
1    Cumings, Mrs. John Bradley (Florence Briggs Th...      0            2   
2                               Heikkinen, Miss. Laina      0            3   
3         Futrelle, Mrs. Jacques Heath (Lily May Peel)      0            4   
4                             Allen, Mr. William Henry      0            5   
..                                                 ...    ...          ...   
886            
```

**Question 1:** Using the describe function, find the mean fare  and store it into `mean_fare`, rounded off to three decimal places. You cannot directly call the .mean() function.

```python
describe = titanic_df.describe()
print(describe)
```

Output:
```
Age        Fare       Parch  PassengerId      Pclass  \
count  891.000000  891.000000  891.000000   891.000000  891.000000   
mean    29.445196   32.204208    0.381594   446.000000    2.308642   
std     13.244896   49.693429    0.806057   257.353842    0.836071   
min      0.420000    0.000000    0.000000     1.000000    1.000000   
25%     22.000000    7.910400    0.000000   223.500000    2.000000   
50%     30.000000   14.454200    0.000000   446.000000    3.000000   
75%     35.500000   31.000000    0.000000   668.500000    3.000000   
max     80.000000  512.329200    6.000000   891.000000    3.000000   

            SibSp    Survived  Family_Size  
count  891.000000  891.000000   891.000000  
mean     0.523008    0.383838     0.904602  
std      1.102743    0.486592     1.613459  
min      0.000000    0.000000     0.000000  
25%      0.000000    0.000000     0.000000  
50%      0.000000    0.000000     0.000000  
75%      1.000000    1.000000     1.000000  
max      8.000000    1.
```

```python

mean_fare = titanic_df["Fare"].describe()['mean']
mean_fare = mean_fare.round(3)



print(f'\nMean Fare is: {mean_fare}')
```

Output:
```
Mean Fare is: 32.204
```

```python
grader.check("Q1")
```

Output:
```
Q1 results: All test cases passed!
```

**Question 2:** Drop the `Cabin` and `Family_Size` column inside the dataframe. `titanic_df` should no longer contain these columns.

```python
titanic_df = titanic_df.drop(['Cabin', 'Family_Size'], axis=1)

```

```python
grader.check("Q2")
```

Output:
```
Q2 results: All test cases passed!
```

```python
titanic_df
```

Output:
```
Age Embarked     Fare  \
0    22.0        S   7.2500   
1    38.0        C  71.2833   
2    26.0        S   7.9250   
3    35.0        S  53.1000   
4    35.0        S   8.0500   
..    ...      ...      ...   
886  27.0        S  13.0000   
887  19.0        S  30.0000   
888  22.0        S  23.4500   
889  26.0        C  30.0000   
890  32.0        Q   7.7500   

                                                  Name  Parch  PassengerId  \
0                              Braund, Mr. Owen Harris      0            1   
1    Cumings, Mrs. John Bradley (Florence Briggs Th...      0            2   
2                               Heikkinen, Miss. Laina      0            3   
3         Futrelle, Mrs. Jacques Heath (Lily May Peel)      0            4   
4                             Allen, Mr. William Henry      0            5   
..                                                 ...    ...          ...   
886                              Montvila, Rev. Juozas      0          887   
887      
```

**Question 3:** What percentage of people aboard the ship were females? Store the result rounded off to three decimal places in `female_percentage`.

```python
total_passengers = len(titanic_df)
female_count = titanic_df['Sex'].value_counts()['female']

female_percentage = round((female_count / total_passengers) * 100, 3)

```

```python
grader.check("Q3")
```

Output:
```
Q3 results: All test cases passed!
```

**Question 4:** Set the `PassengerID` of passengers as the index for this dataset.

```python
titanic_df.columns
```

Output:
```
Index([' Age', 'Embarked', 'Fare', 'Name', 'Parch', 'PassengerId', 'Pclass',
       'Sex', 'SibSp', 'Survived', 'Ticket', 'Title'],
      dtype='object')
```

```python
titanic_df = titanic_df.set_index('PassengerId', drop= True)

```

```python
titanic_df.head(5)
```

Output:
```
Age Embarked     Fare  \
PassengerId                           
1            22.0        S   7.2500   
2            38.0        C  71.2833   
3            26.0        S   7.9250   
4            35.0        S  53.1000   
5            35.0        S   8.0500   

                                                          Name  Parch  Pclass  \
PassengerId                                                                     
1                                      Braund, Mr. Owen Harris      0       3   
2            Cumings, Mrs. John Bradley (Florence Briggs Th...      0       1   
3                                       Heikkinen, Miss. Laina      0       3   
4                 Futrelle, Mrs. Jacques Heath (Lily May Peel)      0       1   
5                                     Allen, Mr. William Henry      0       3   

                Sex  SibSp  Survived            Ticket Title  
PassengerId                                                   
1              male      1         0         A
```

```python
grader.check("Q4")
```

Output:
```
Q4 results: All test cases passed!
```

**Question 5:** Convert all column names to lowercase and remove any leading and trailing whitespaces. Simultaneously, replace any non-leading and non-trailing spaces within the column names with an underscore. You are not allowed to hardcode the column names.

```python
titanic_df.columns = titanic_df.columns.str.lower().str.strip()


```

```python
titanic_df.columns
```

Output:
```
Index(['age', 'embarked', 'fare', 'name', 'parch', 'pclass', 'sex', 'sibsp',
       'survived', 'ticket', 'title'],
      dtype='object')
```

```python
grader.check("Q5")
```

Output:
```
Q5 results: All test cases passed!
```

**Question 6:** Find the female passenger that had the highest fare and store the result in `highest_fare`. Furthermore, find the male passenger that had the lowest non-zero fare and store it in `lowest_fare`. <br>
Note: You are not meant to store the value only but the entire row entry.

```python
titanic_df.columns
```

Output:
```
Index(['age', 'embarked', 'fare', 'name', 'parch', 'pclass', 'sex', 'sibsp',
       'survived', 'ticket', 'title'],
      dtype='object')
```

```python
#  female passenger with the highest fare
highest_fare = titanic_df[titanic_df['sex'] == 'female'].loc[[titanic_df[titanic_df['sex'] == 'female']['fare'].idxmax()]]

# male passenger with the lowest non-zero fare
lowest_fare = titanic_df[(titanic_df['sex'] == 'male') & (titanic_df['fare'] > 0)].loc[[titanic_df[(titanic_df['sex'] == 'male') & (titanic_df['fare'] > 0)]['fare'].idxmin()]]

# Print
print(highest_fare)
print(lowest_fare)

```

Output:
```
age embarked      fare              name  parch  pclass     sex  \
PassengerId                                                                     
259          35.0        C  512.3292  Ward, Miss. Anna      0       1  female   

             sibsp  survived    ticket title  
PassengerId                                   
259              0         1  PC 17755  Miss  
              age embarked    fare                 name  parch  pclass   sex  \
PassengerId                                                                    
379          20.0        C  4.0125  Betros, Mr. Tannous      0       3  male   

             sibsp  survived ticket title  
PassengerId                                
379              0         0   2648    Mr
```

```python
grader.check("Q6")
```

Output:
```
Q6 results: All test cases passed!
```

**Question 7:** Write a function `CalculateSurvival(data)` that takes in a dataframe and returns the survival percentage rounded off to three decimal places.

```python
def CalculateSurvivalRate(data):
    survivors = data['survived'].sum()
    count = data['survived'].count()
    survival_percentage = (survivors / count) * 100
    
    return round(survival_percentage,3)
```

```python
grader.check("Q7")
```

Output:
```
Q7 results: All test cases passed!
```

**Question 8A:** Add a column called `family` to the dataframe that shows true for a passenger whose family members (siblings, spouses, parents, or children) were on board and false otherwise.

```python
unique_strings = titanic_df['sibsp'].unique()

# Display the unique strings
print(unique_strings)
```

Output:
```
[1 0 3 4 2 5 8]
```

```python
titanic_df.head(8)
```

Output:
```
age embarked     fare  \
PassengerId                           
1            22.0        S   7.2500   
2            38.0        C  71.2833   
3            26.0        S   7.9250   
4            35.0        S  53.1000   
5            35.0        S   8.0500   
6            30.0        Q   8.4583   
7            54.0        S  51.8625   
8             2.0        S  21.0750   

                                                          name  parch  pclass  \
PassengerId                                                                     
1                                      Braund, Mr. Owen Harris      0       3   
2            Cumings, Mrs. John Bradley (Florence Briggs Th...      0       1   
3                                       Heikkinen, Miss. Laina      0       3   
4                 Futrelle, Mrs. Jacques Heath (Lily May Peel)      0       1   
5                                     Allen, Mr. William Henry      0       3   
6                                             Moran, Mr.
```

```python
titanic_df['family'] = (titanic_df['sibsp'] > 0) | (titanic_df['parch'] > 0)
```

```python
titanic_df.head()
```

Output:
```
age embarked     fare  \
PassengerId                           
1            22.0        S   7.2500   
2            38.0        C  71.2833   
3            26.0        S   7.9250   
4            35.0        S  53.1000   
5            35.0        S   8.0500   

                                                          name  parch  pclass  \
PassengerId                                                                     
1                                      Braund, Mr. Owen Harris      0       3   
2            Cumings, Mrs. John Bradley (Florence Briggs Th...      0       1   
3                                       Heikkinen, Miss. Laina      0       3   
4                 Futrelle, Mrs. Jacques Heath (Lily May Peel)      0       1   
5                                     Allen, Mr. William Henry      0       3   

                sex  sibsp  survived            ticket title  family  
PassengerId                                                           
1              male      1    
```

```python
titanic_df['family'].value_counts()[True] 
```

Output:
```
354
```

```python
grader.check("Q8A")
```

Output:
```
Q8A results: All test cases passed!
```

**Question 8B:** Using the `CalculateSurvivalRate` function, find the survival percentage of the passengers who had family on board (siblings, spouses, children or parents) and store the result in `survival_rate_with_family`.

```python
# Filter  passengers with family on board
family_passengers = titanic_df[titanic_df['family']]

# Calculate the total number of family passengers
total_family_passengers = len(family_passengers)

# Calculate the number of survivors among family passengers
survivors_with_family = family_passengers['survived'].sum()

# Calculate the survival rate as a percentage
if total_family_passengers > 0: 
    survival_rate_with_family = (survivors_with_family / total_family_passengers) * 100
else:
    survival_rate_with_family = 0  

# Round off
survival_rate_with_family = round(survival_rate_with_family, 3)


print(survival_rate_with_family)

```

Output:
```
50.565
```

```python
grader.check("Q8B")
```

Output:
```
Q8B results: All test cases passed!
```

**Question 8C:** Using the `CalculateSurvivalRate` function, find the survival percentage of the passengers who did not have any family on board (no siblings, spouses, children and parents) and store the result in `survival_rate_without_family`.

```python
# Filter passengers without family on board
no_family_passengers = titanic_df[titanic_df['family'] == False]

# Calculate the total number of no family 
total_no_family_passengers = len(no_family_passengers)

# Calculate the number of survivors among no family 
survivors_without_family = no_family_passengers['survived'].sum()

# Calculate the survival rate as a percentage
if total_no_family_passengers > 0:  # Avoid division by zero
    survival_rate_without_family = (survivors_without_family / total_no_family_passengers) * 100
else:
    survival_rate_without_family = 0  # If there are no no family passengers

# Round off
survival_rate_without_family = round(survival_rate_without_family, 3)


print(survival_rate_without_family)

```

Output:
```
30.354
```

```python
grader.check("Q8C")
```

Output:
```
Q8C results: All test cases passed!
```

**Question 8D:** Based on your findings above, were the passengers who had family more likely to survive than those who did not? Answer as 'Yes' or 'No' and store it in `ans`.

```python
ans = 'Yes'
```

```python
grader.check("Q8D")
```

Output:
```
Q8D results: All test cases passed!
```

**Question 9:** Use `groupby` to find the survival percentage of males and females and store them in `male_survival_percentage` and `female_survival_percentage` respectively. Round off your answers to three decimal places. Not using `groupby` will result in a zero for this question.

```python
survival_rates = titanic_df.groupby('sex')['survived'].mean() * 100
```

```python
male_survival_percentage = round(survival_rates['male'],3)
female_surival_percentage = round(survival_rates['female'],3)

```

```python
grader.check("Q9")
```

Output:
```
Q9 results: All test cases passed!
```

**Question 10:** Using the pandas iloc function, extract the rows with passengers aged between 5 to 20 and fare above 20. Store this result in `ilocresult`. 

```python
ilocresult = titanic_df[(titanic_df['age'] >= 5) & (titanic_df['age'] < 20) & (titanic_df['fare'] > 20)]
ilocresult = ilocresult[['age','fare']]
```

```python
ilocresult.shape
```

Output:
```
(53, 2)
```

```python
grader.check("Q10")
```

Output:
```
Q10 results: All test cases passed!
```

**Question 11A:** Replace the integer values in the `Pclass` columns with corresponding string data types. For example, replace 1 with 'Upper', 2 with 'Middle' and 3 with 'Lower'.

```python
titanic_df['pclass'] = titanic_df['pclass'].replace({1: 'Upper', 2: 'Middle', 3: 'Lower'})
```

```python
grader.check("Q11A")
```

Output:
```
Q11A results: All test cases passed!
```

**Question 11B:** Using a pivot table, represent the male and female passengers that survived in each of the three passenger classes. Store the result in `class_count`.

```python
class_count = titanic_df.pivot_table(
    index='pclass',          # 
    columns='sex',          #  gender as the columns
    values='survived',      # Use the 'Survived' column for aggregation
    aggfunc='sum'           # Aggregate by summing the survived values (1 for survived, 0 for not)
)
```

```python
class_count
```

Output:
```
sex     female  male
pclass              
Lower       72    47
Middle      70    17
Upper       91    45
```

```python
grader.check("Q11B")
```

Output:
```
Q11B results: All test cases passed!
```

**Well done on completing the first part of your assignment. Take a break and move onto the next!**

Tests for Q0 (points: 1)
```
>>> assert titanic_df.shape == (891, 14)
```


Tests for Q1 (points: 3)
```
>>> assert mean_fare == 32.204
```


Tests for Q10 (points: 3)
```
>>> assert ilocresult.shape == (53, 2)
>>> assert ilocresult['age'].median() == 15.0
>>> assert ilocresult['fare'].median() == 34.375
```


Tests for Q11A (points: 2)
```
>>> assert titanic_df['pclass'].value_counts()['Upper'] == 216
>>> assert titanic_df['pclass'].value_counts()['Lower'] == 491
>>> assert titanic_df['pclass'].value_counts()['Middle'] == 184
```


Tests for Q11B (points: 4)
```
>>> assert list(class_count['female']) == [72, 70, 91]
>>> assert list(class_count['male']) == [47, 17, 45]
```


Tests for Q2 (points: 3)
```
>>> assert titanic_df.shape == (891, 12)
>>> assert 'Cabin' not in titanic_df.columns
>>> assert 'Family_Size' not in titanic_df.columns
```


Tests for Q3 (points: 3)
```
>>> assert female_percentage == 35.241
```


Tests for Q4 (points: 2)
```
>>> assert 'PassengerId' not in titanic_df.columns
>>> assert titanic_df.index.name == 'PassengerId'
>>> assert titanic_df.shape == (891, 11)
```


Tests for Q5 (points: 3)
```
>>> assert list(titanic_df.columns) == ['age', 'embarked', 'fare', 'name', 'parch', 'pclass', 'sex', 'sibsp', 'survived', 'ticket', 'title']
```


Tests for Q6 (points: None)
```
>>> assert highest_fare.index[0] == 259
>>> assert highest_fare['fare'].values[0] == 512.3292
>>> assert lowest_fare['fare'].values[0] == 4.0125
>>> assert lowest_fare.index[0] == 379
```


Tests for Q7 (points: 4)
```
>>> assert CalculateSurvivalRate(titanic_df.loc[(titanic_df['sibsp'] > 0) | (titanic_df['parch'] > 0)]) == 50.565
>>> assert CalculateSurvivalRate(titanic_df.loc[(titanic_df['sibsp'] == 0) & (titanic_df['parch'] == 0)]) == 30.354
```


Tests for Q8A (points: 2)
```
>>> assert titanic_df['family'].value_counts()[True] == 354
>>> assert titanic_df['family'].value_counts()[False] == 537
```


Tests for Q8B (points: 2)
```
>>> assert survival_rate_with_family == 50.565
```


Tests for Q8C (points: 2)
```
>>> assert survival_rate_without_family == 30.354
```


Tests for Q8D (points: 2)
```
>>> assert ans.strip().lower() == 'yes'
```


Tests for Q9 (points: 3)
```
>>> assert male_survival_percentage == 18.891
>>> assert female_surival_percentage == 74.204
```
